## Using Internet Movie Database

Using the Internet Movie Database file `titles.csv.gz`, using `pandas` find the how many movies were
made every year from `1900` to `2016`. The data files are available at the following links:

[titles](https://storage.googleapis.com/class-notes-181217.appspot.com/titles.csv.gz)
[cast](https://storage.googleapis.com/class-notes-181217.appspot.com/cast.csv.gz)

<!-- grade_id: cell-0a22d8809d9aab94, points: 0, grade: false, solution: true -->
```python
import pandas as pd

file_ip = pd.read_table("https://storage.googleapis.com/class-notes-181217.appspot.com/titles.csv.gz",compression='gzip', delimiter = ',') 
counter = 0
title_year = file_ip['year']
for i in range(len(title_year)):
    if ( 1900 <= title_year[i] ) and ( title_year[i] <= 2016):
        counter +=1

answer = counter

```

<!-- grade_id: cell-a437770e91bcdfa0, points: 10, grade: true, solution: false -->
```python
### BEGIN TESTS
assert answer == 209649
### END TESTS

```

Using the cast IMDB file `cast.csv.gz`, determine how many movies the actor *John Wayne* was in (under that name).

<!-- grade_id: cell-3fd954514f1e2d94, points: 0, grade: false, solution: true -->
```python

import pandas as pd
cast = pd.read_table("https://storage.googleapis.com/class-notes-181217.appspot.com/cast.csv.gz",compression='gzip', delimiter = ',')
cast_counter = cast["name"].value_counts()
answer = cast_counter["John Wayne"]

```

<!-- grade_id: cell-799f9c7d314da863, points: 10, grade: true, solution: false -->
```python
### BEGIN TESTS
assert answer == 179
### END TESTS

```

In how many movies did John Wayne play a character named *John* (not Johnny, Jonny, etc.)? This includes *Captain John* and others like this.

<!-- grade_id: cell-3904b4d8eeaf5967, points: 0, grade: false, solution: true -->
```python
import pandas as pd 
cast = pd.read_table("https://storage.googleapis.com/class-notes-181217.appspot.com/cast.csv.gz",compression='gzip', delimiter = ',')
cast_character_numbered = cast["character"].value_counts()
John_Wayne_df = cast[cast.name == "John Wayne"]
answer=John_Wayne_df.character.str.contains(r'\bJohn\b').sum()

```

<!-- grade_id: cell-9f214890c9fc7de2, points: 10, grade: true, solution: false -->
```python
### BEGIN  TESTS
assert answer == 38
### END  TESTS

```

```python

```

```python

```